Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 Math

What will be the output of the following lines of code?  (Yes, I know that you can easily copy and paste this into another notebook and run it.  Please try to reason through it!)

```python
temp_c = 100
temp_f = temp_c * 9//5 + 32
temp_f
```

> A) 212
>
> B) 132
>
> C) 32
> 
> D) 98.6

<!-- grade_id: cell-7eff6a222bd88fa0, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: cell-71d7d61ddc7cd8a2, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 Higher-Level Math

Which code represents the proper way to represent the mathematical expression $pi^2$ ?

> A) `pi^2`
>
> B) `pi x pi`
>
> C) `pi ** 2`
>
> D| `pi | 2`

<!-- grade_id: cell-5421f338f0f96d36, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "C"
    
    return answer
```

<!-- grade_id: cell-db4cf8b387a07b24, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 Combining Data

Which of the following is not a valid operation?

> A) `'This is a ' + "complete sentence."`
>
> B) `'She is ' + 93 + ' years old.'`
>
> C) `"Repeat this! " * 4`
>
> D) `3.14 * 7`

<!-- grade_id: cell-afd3553a028a91fb, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: cell-d63f3eb2259c6cc9, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```