Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - True or False: In a Pandas dataframe, each row has to be unique.  That is, there cannot be any duplicate rows in a Pandas data frame.

> True or False


<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = False
    
    return answer
```

<!-- grade_id: q1-solution, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == bool
```

## #2 - Which of the following will NOT give you (at least) a unique list of values for the `state` column in the data frame `data`?

> A) `data.unique()`
> 
> B) `data.groupby('State').count()`
>
> C) `data.State.value_counts()`
>
> D) `data['State'].unique()`


<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: q2-solution, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - Which of the following data sources can Pandas load directly?  

If there is more than one answer, put your answers in a list.

> A) CSV from the Internet
>
> B) JSON in a file
>
> C) Microsoft Access Database
>
> D) Excel


<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = ["A","C"]
    
    return answer
```

<!-- grade_id: q3-solution, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == list
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-2-4c53244c44fd> in <module>
----> 1 assert type(q03()) == list

AssertionError:
```

```python

```