Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - In this class, XML Stands for:

> A) eXtensible Markup Language
>
> B) The 10th version of the Markup Language
>
> C) eXhange Markup Language
>
> D) X-Men Legends


<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: q1-solution, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - Which of the following is not a type of data in an XML document?

> A) tags
>
> B) attributes
>
> C) text
>
> D) list


<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "D"
    
    return answer
```

<!-- grade_id: q2-solution, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - True or False, each element in an XML document can only have one child with a specific tag.

For example, is the following XML illegal because there are two child elements with the tag "kid":

```xml
<family>
  <parent role="dad">Paul Boal</parent>
  <kid role="daughter">Eleanor</kid>
  <kid role="daughter">Adelaide</kid>
</family>
```

Answer `False` if you think the above XML is valid.  

Answer `True` if you think the above XML is invalid.




<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = False
    
    return answer
```

<!-- grade_id: q3-solution, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == bool
```