Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - What does the keyword `def` mean before a function name?

> A) _define_ - You are saying that this is the function definition
>
> B) _default_ - This is the default work the function should do
>
> C) _def_ - Doesn't stand for anything, it's just what you put there
>
> D) _defer_ - Meaning don't do anything right now

<!-- grade_id: cell-d789a32349f1885c, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: cell-ebb0bb10a19c592a, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - What is the minimum number of parameters that you have to provide for this function to work?

```python
def my_function (fname='John', lname='Doe', gender='M'):
    ...
```

> A) 3
>
> B) 2
>
> C) 1
>
> D) 0

<!-- grade_id: cell-7ec686102170454b, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "D"
    
    return answer
```

<!-- grade_id: cell-0592dde608b988be, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - Which of the following lines of code does not have any functions in it?


> A) `value = math.log(23,2)`
>
> B) `pi * (radius ** 2)`
>
> C) `[ x for x in range(10) ]`
>
> D) `print('I am a good student')`

<!-- grade_id: cell-46ee97a89206f3d4, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: cell-be350b2fc623fd53, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```