Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Setup Jupyter to talk to Github.com

Now that you've gotten signed into Jupyter, we're going to have to get let them talk in the background without prompting you for your password over and over.  To so this, we use something called a private/public key pair.  This series of instructions will effectively created a lock and key pair and let you install the lock (public key) on your Github.com account.  Then anyone with the key (private key) can access your Github.com account using it - and you'll simply keep this key stashed privately on the https://jupyter.slucor.net.

## Step 1: Creating the keys

```python
%%bash
cd ~

if test -e ~/.ssh/id_rsa; then
  echo "Skipped ssh-keygen since you already have a key pair created." 1>&2
else
  ssh-keygen -t rsa -N "" -f ~/.ssh/id_rsa -q
  echo "Generating a new SSH key pair to use with github.com" 1>&2
fi

cat ~/.ssh/id_rsa.pub
```

Output:
```
ssh-rsa AAAAB3NzaC1yc2EAAAADAQABAAABAQCvKQMZjsnliEa26UG3aaiNo8gZzqE58C0nW52FMliJWzKv9L1cwQQpt1JXcuWsPB1e3NlgFMENhgLGZKqtBtHAHMzRAtsWgcaBefarUZS/SatsiFpznYhMrYWKPYGTE3lZmmlHyVGwiMxisaPUZ5UzHVSoJwklWcGw7MT9zrPlb1GIce8CsPQDdvgF8LBo4KKl1vQCXW18rJ+TtJzjWW/hPj0v6vsFPU0s/Hek5fTNFjZaHOD913SmmE0PxgGrXM7DxYMp3vIQIGX6UiUycfjwQQhTr/59gwH7dmDesUx14DpVDk0OUqKYeFJnwXJwLStxs8GYrmk2HbEe/B9KzfVn eeholcomb@ip-172-31-48-250.ec2.internal
Generating a new SSH key pair to use with github.com
```

<!-- grade_id: verify-keys-match, points: 1, grade: true, solution: false -->
```python
from Crypto.PublicKey import RSA
import os

with open(os.path.expanduser('~/.ssh/id_rsa')) as f:
    private = RSA.import_key(f.read())

with open(os.path.expanduser('~/.ssh/id_rsa.pub')) as f:
    public = f.read()
    
public1 = public[0:public.find(' ',10)]
public2 = private.export_key('OpenSSH').decode('ASCII')

assert(public1 == public2)

```

## Step 2: Put the public key in github

Copy to your clipboard he text that you have above, with the public key.

1. Go login to https://github.com
2. Click on your profile icon in the upper right corner of the window.
3. Choose **Settings**
4. Choose **SSH and GPG keys** on the right-hand navigation area
5. Click the **New SSH key** button
6. Put `HDS5210 Jupyter` in the **Title** and the information you copied from above into the **Key** field.
7. Click the **Add SSH key** button


