Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - What is the output from the code below?

```python
dxs = [['COPD','CHF','MI'],['DVT','DBS'],['GI','CNC']]
len(dxs[1])
```

> A) _Error_
>
> B) `3`
>
> C) `2`
>
> D) _None_

<!-- grade_id: cell-d789a32349f1885c, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "C"
    
    return answer
```

<!-- grade_id: cell-ebb0bb10a19c592a, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - What will the resulting value of `vals` be?

```python
vals = [3.14, 1.41]
vals.append([2.72])
vals = vals + ['i']
```

> A) `[3.14, 1.41, 2.72, ['i']]`
>
> B) `[3.14, 1.41, [2.72], 'i']`
>
> C) `[3.14, 1.41, [2.72], ['i']]`
>
> D) `[3.14, 1.41, [2.72, 'i']`

<!-- grade_id: cell-7ec686102170454b, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: cell-0592dde608b988be, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - Which of the following statements is **not** correct?


> A) Lists can have other lists inside them
>
> B) When you create a list, you must specify how big it will be
>
> C) A list is an ordered set of objects
>
> D) list.sort() will sort the list in place without creating a copy

<!-- grade_id: cell-46ee97a89206f3d4, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: cell-be350b2fc623fd53, points: 1, grade: true, solution: false -->
```python

assert type(q03()) == str
```

```python

```