Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 Code Blocks

Why is some code in Python indented from other lines of code?
> A) To specify that line of code should only be executed if the previous line is True
>
> B) The programmer can choose when to indent code to make it easier to read
>
> C) To block code apart to designate it should only be executed at certain times
>
> D) Indented lines are treated like comments and not executed

<!-- grade_id: cell-e3bfb833d65e06cf, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "C"
    
    return answer
```

<!-- grade_id: cell-97b796d3d10e8bc9, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 Python Versions

Which statements about Python 2 and Python 3 are true?
> A) Python version 2 and version 3 have completely compatible syntax
>
> B) Any code that runs in Python version 3 will run in version 2
> 
> C) Python version 3 and version 2 packages are compatible
> 
> D) Python version 2 and version 3 have some differences, but are mostly the same

<!-- grade_id: cell-4343e4e3597378de, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "D"
    
    return answer
```

<!-- grade_id: cell-c44c3d14e219403e, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 Python Syntax

Which other common programming languages does Python look like?
> A) C++
>
> B) Java
> 
> C) JavaScript
>
> D) None of these

<!-- grade_id: cell-77203bf463dcb9cb, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "D"
    
    return answer
```

<!-- grade_id: cell-719b9ddb4e7b8277, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```