Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - If you accidentally run an infine while loop in Jupyter:

> A) Close the browser window and it will automatically stop.
>
> B) Interrupt or restart the kernel.
>
> C) You will overheat the server and burn it up.
>
> D) Slack the professor so that he can stop the loop for you.

<!-- grade_id: cell-d789a32349f1885c, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: cell-ebb0bb10a19c592a, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - What will the following code print to the screen?

```python
counter = 1

while (counter < 10):
    print(counter)
    break
```

> A) `1`
>
> B) `1,2,3,4,5,6,7,8,9`
>
> C) `1,2,3,4,5,6,7,8,9,10`
>
> D) _nothing_

<!-- grade_id: cell-7ec686102170454b, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: cell-0592dde608b988be, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - What will the follow code do?

```python
status = True
while (status):
    status =  False
while (status):
    status = True
    print("The status is now true")
```

> A) Only the first line of code will execute
>
> B) The code will print 'The status is now true'
>
> C) The code will all run but nothing will print
>
> D) None of these

<!-- grade_id: cell-46ee97a89206f3d4, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "C"
    
    return answer
```

<!-- grade_id: cell-be350b2fc623fd53, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```