Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Descriptive Stats

For the questions below, we're going to refer exclusively to an HHCAPS survey data set that is available in `/data/hhcaps.csv`.  Use whatever commands you want to calculate the information required to get to the answer.


Put your solution as the last line right before the test cell, using the form shown below:

```
answer = ...
```



```python
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv')
```

```python
hhcaps.head()
```

Output:
```
State  CMS Certification Number (CCN)*  \
0    AL                            17000   
1    AL                            17008   
2    AL                            17009   
3    AL                            17013   
4    AL                            17014   

                                       Provider Name  \
0  BUREAU OF HOME & COMMUNITY SERVICES           ...   
1  JEFFERSON COUNTY HOME CARE                    ...   
2  ALACARE HOME HEALTH & HOSPICE                 ...   
3  GENTIVA HEALTH SERVICES                       ...   
4  AMEDISYS HOME HEALTH                          ...   

                                             Address                  City  \
0  201 MONROE STREET, THE RSA TOWER,  SUITE 1200 ...  MONTGOMERY             
1  2201 ARLINGTON AVENUE                         ...  BESSEMER               
2  2970 LORNA ROAD                               ...  BIRMINGHAM             
3  1239 RUCKER BLVD                              ...  ENTERPRISE             
4  68278 M
```

## #01 - 

How many columns does this file contain?

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

# YOUR CODE HERE
answer = len(hhcaps.columns)
```

<!-- grade_id: q01-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)


```

```python
answer
```

Output:
```
39
```

## #02 - 

How many different values for State are there in this file?

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

# YOUR CODE HERE
answer = hhcaps['State'].nunique()
```

<!-- grade_id: q02-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
```

```python
answer
```

Output:
```
55
```

## #03 - 

Which of those State values has the highest frequency of occurence?

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
answer = ''

# YOUR CODE HERE
answer = hhcaps.groupby('State').size().idxmax()
```

<!-- grade_id: q03-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
```

```python
answer
```

Output:
```
'TX'
```

## #04 - 

Which of those State values has the best average performance on the `Star Rating for health team communicated well with them` score?

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
answer = ''

# YOUR CODE HERE
answer = hhcaps.groupby('State').mean()['Star Rating for health team communicated well with them'].idxmax()
```

<!-- grade_id: q04-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
```

```python
answer
```

Output:
```
'ME'
```

## #05 - 

What was the average score on `Star Rating for how patients rated overall care from agency` for providers listed as having a `Type of Ownership` of `Hospital Based Program`

<!-- grade_id: q05-solution, points: 0, grade: false, solution: true -->
```python
answer = 0.0

# YOUR CODE HERE
answer = hhcaps[hhcaps['Type of Ownership'] == 'Hospital Based Program']['Star Rating for how patients rated overall care from agency'].mean()
```

<!-- grade_id: q05-test, points: 2, grade: true, solution: false -->
```python
import numpy
assert(type(answer) == float or type(answer) == numpy.float64)
```

```python
answer
```

Output:
```
3.6925207756232687
```

## #06 - 

There is one column in the file that has the same value on every row.  What is the name of that column?

<!-- grade_id: q06-solution, points: 0, grade: false, solution: true -->
```python
answer = ''

# YOUR CODE HERE
answer = hhcaps.apply(lambda x: x.nunique()).idxmin()
```

<!-- grade_id: q06-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
```

```python
answer
```

Output:
```
'Offers Nursing Care Services'
```

## #07 - 

There are six (6) columns in the file that indicate (True or False) if the provider offers certain services.  Those column names all start with `Offers...`.  How many facilities offer **everything** that the survey was interested in asking about?  That is, how many rows have True in all six of those columns?

<!-- grade_id: q07-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

# YOUR CODE HERE
answer = hhcaps[hhcaps[hhcaps.columns[hhcaps.columns.str.match('Offers')]].all(axis='columns')].shape[0]
```

<!-- grade_id: q07-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
```

```python
answer
```

Output:
```
9029
```

## #08 - 

The survey results report `Number of completed Surveys` as one of the metrics.  What is the `median` of the number of completed surveys per facility?

Note that some rows don't have a valid number in them for the number of completed surveys.


<!-- grade_id: q08-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

# YOUR CODE HERE
answer = hhcaps[hhcaps['Number of completed Surveys'].apply(lambda x: x.isnumeric())]['Number of completed Surveys'].median()
```

<!-- grade_id: q08-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == float or type(answer) == numpy.float64
)
```

```python
answer
```

Output:
```
84.0
```

## #09 - 

How many providers are there from St. Louis, Missouri?

*Note that the City columnd may have trailing spaces that need to be accomodated or cleaned up*

<!-- grade_id: q09-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

# YOUR CODE HERE
hhcaps['City'] = hhcaps['City'].str.strip()

answer = hhcaps[hhcaps['City'] == 'SAINT LOUIS']['Provider Name'].nunique()
```

<!-- grade_id: q09-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
```

```python
answer
```

Output:
```
23
```

# #10 - 

Which state got the highest percentage of 'top box' scorings, using the `Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)` field.

Ignore those rows where the data is **Not Available**

```python

```

<!-- grade_id: q10-solution, points: 0, grade: false, solution: true -->
```python
answer = 0
field = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'

# YOUR CODE HERE
rating_hhcaps = hhcaps[hhcaps[field].apply(lambda x: x.isnumeric())]
rating_hhcaps = rating_hhcaps[rating_hhcaps['Number of completed Surveys'].apply(lambda x: x.isnumeric())]
rating_hhcaps['calc'] = rating_hhcaps[field].astype('int') * rating_hhcaps['Number of completed Surveys'].astype('int')
rating_hhcaps['Number of completed Surveys'] = rating_hhcaps['Number of completed Surveys'].astype('int')

rating = rating_hhcaps.groupby(['State']).agg({'Number of completed Surveys':'sum','calc':'sum'})

answer = (rating['calc'] / rating['Number of completed Surveys']).idxmax()
```

<!-- grade_id: q10-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
```

```python
answer
```

Output:
```
'MS'
```

```python

```