Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - Which is true about docstrings?

> A) They are stored in Python as just a long string inside the function definition.
>
> B) They can contain tests to help you demonstrate your code works correctly.
>
> C) They tell Python what to say when someone asks for `help()` on your function.
>
> D) All of these

<!-- grade_id: cell-d789a32349f1885c, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "D"
    
    return answer
```

<!-- grade_id: cell-ebb0bb10a19c592a, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - Explain in your own words what "Test Driven Development" means


<!-- grade_id: cell-7ec686102170454b, points: 1, grade: true, solution: true -->
```python
def q02():

# YOUR CODE HERE
    answer = "Test Driven Development is when you test a function you've written, and do what is required to fix it if it fails. After it's fixed, you add another test and if it fails, you fix it again. I think this will help improve and widen the scope of a function to minimize hard-coding and see if it will work with different types of data relevant to the intention of the function."
    
    return answer
```

<!-- grade_id: cell-0592dde608b988be, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - True or False: From now on you should include docstrings in the _applied_ homework assignments?

Type `answer = True` or `answer = False` as your response

<!-- grade_id: cell-46ee97a89206f3d4, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = True
    
    return answer
```

<!-- grade_id: cell-be350b2fc623fd53, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == bool
```