Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - When a Python program opens a file to read from it, no other programs can read from the same file at the same time.

> True or False


<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = False
    
    return answer
```

<!-- grade_id: q1-solution, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == bool
```

## #2 - What is the advantage of opening files in Python using a `with` block?

> A) It ensures the file is closed as soon as you're done using it
>
> B) It runs faster than not using the `with` block
>
> C) It automatically loads the whole file into memory so that it can be accessed faster
>
> D) Actually, it is the only way to do it


<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: q2-solution, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - **ALL** files are written using plain text.  Computers just use a lanugage like JSON or XML to explain what the file contents mean.

> A) Many files only contain text, but many file formats do not use text
>
> B) All files contain only text, it's the just the text contents that explain the data to the computer
>
> C) All files contain only text in either JSON, XML, or CSV.
>
> D) Even if files contain only text, there still has to be a special code at the beginning to explain that this file will contain only text.


<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: q3-solution, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```