Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 Built-in Types

Which of the following are NOT standard built-in data types in Python.  Put your answer inside brackets as shown here: `answer = ["X", "Y", "Z"]` 

> A) bool
>
> B) datetime
>
> C) decimal
>
> D) float
>
> E) int
>
> F) timestamp
>
> G) None


<!-- grade_id: cell-80dc9ca5c675bbde, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # answer = ["X","Y","Z"]
    # YOUR CODE HERE
    answer = ["B","C","F"]
    
    return answer
```

<!-- grade_id: cell-a08ca8cf94a4e835, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == list
```

## #2 Strings

Which is the proper way to create a multi-line string in a Python program?

> A) Place `#` characters at the beginning of each line of the string
>
> B) Start the string with a `'` or `"` and then indent all of the subsquent lines until you place a `'` or `"` at the end of the multi-line string
>
> C) Place triple-quotes (`'''` or `"""`) on the first line and last line of the multi-line string
>
> D) You can't create multi-line strings in Python


<!-- grade_id: cell-2946b0bb4ae276e6, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # Use the format answer = "X" again
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: cell-250a2cc17824ea0c, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 Variable Types

Which statement is **not** true about the following segment of code:
```
var = "My first program"
var = 3.14
var = var + " is pi"
```

> A) `var` is a very bad name for a variable
>
> B) The code will fail on the second line because `var` is already a string and can't be changed to a number
>
> C) The code will fail on the last line because `var` is a number and you can't add numbers and strings
>
> D) `3.14` has the type of `float`


<!-- grade_id: cell-b2a1d9e47b8fb90c, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: cell-fac4de98eef8e85b, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```