Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - What is another name for the inline if statement?

> A) Ternary operator
>
> B) Conditional
>
> C) Three-part if
>
> D) None of these

<!-- grade_id: cell-d789a32349f1885c, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: cell-ebb0bb10a19c592a, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - What does this code output?

```python
age = 67
gender = 'male'

print(
    ('young ' if age < 18 else ('retired ' if age > 65 else '')) + 
    ('woman' if gender == 'woman' else 'man')
)
```

> A) `man`
>
> B) `retired man`
>
> C) `retired`
>
> D) _nothing_

<!-- grade_id: cell-7ec686102170454b, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: cell-0592dde608b988be, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - Which of the following statements is not correct?


> A) The inline if can be used with as many `elif` clauses as needed
>
> B) Inline if can be used just about anywhere in your code 
>
> C) Inline if can return any kind of object
>
> D) Inline if is a useful way to streamline code

<!-- grade_id: cell-46ee97a89206f3d4, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: cell-be350b2fc623fd53, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```