Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

## #1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
def longest(L):
    """
    This function will return the first position at which the longest string exists.
    The function requires only one argument, and it must be a list of strings.
    """
    pos = []
    
    for position, string in enumerate(L):
        if string == max(L,key=len):
            pos = position
            
    return pos
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

## #2 Exception Handling

Write a function called `safe_average(L)` to calculate the average of a list of values.  If any of the values in the list are not valid numbers, then the function should simply ignore them and exclude from the average.  If none of the values are valid numbers, then the function should return None.

Example:

```
>>> safe_average([100,200,'three'])
150

>>> safe_average(['one','two','three'])
None
```

<!-- grade_id: exception-answer, points: 0, grade: false, solution: true -->
```python
from statistics import mean

def safe_average(L):
    """
    This function will return the mean for a list of items.
    It will ignore all values except for those that are integers.
    """
    numbers = []
    
    if (all(type(i) != int for i in L)):
        return None
    else:
        L.remove(L[L != int])
        average = mean(L)
    
    return average
```

```python
list = ['one','two','three']

list[list != int or float]
```

```python
all(type(i) is int for i in list)
```

```python
safe_average([1,2,3])

```

<!-- grade_id: exception-tests, points: 10, grade: true, solution: false -->
```python
assert safe_average([1,2,3]) == 2
assert safe_average([1,'two',3]) == 2
assert safe_average(['one','two','three']) == None
```

```python

```