Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Week 4 Exercises

_McKinney 3.1_

**In the exercises below, don't worry about the `def` and other things around the code you need to write.  If there is a skeleton of code already there, just put your code where it says "your code here"**

## #1 Convert a list to a dictionary

Take the list of lists below and convert it into a dictionary where the entries in the dictionary are keyed off of whatever values are in column 0, and the dictionary values are the list of all the values that appear with that key from the input list.

<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
def list_to_dict():

    names = [
             ['Boal', 'Paul'],
             ['Duck', 'Donald'],
             ['Duck', 'Daisy'],
             ['Boal', 'Ada'],
             ['Boal', 'Teddy'],
             ['Westhus', 'Eric']
            ]

    # To help you, I've provided a blank dictionary named `families`.
    families = {}
    
    # Your code should loop through `names` and do something to build
    # up `families` to be the output dictionary shown in the test
    # cells below.

    # YOUR CODE HERE
    for name in names:
        if name[0] in families:
            families[name[0]].append(name[1])
        else:
            families[name[0]] = [name[1]]

    return families

```

```python
list_to_dict()
```

Output:
```
{'Boal': ['Paul', 'Ada', 'Teddy'],
 'Duck': ['Donald', 'Daisy'],
 'Westhus': ['Eric']}
```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
answer = {
 'Boal':    ['Paul', 'Ada', 'Teddy'],
 'Duck':    ['Donald', 'Daisy'],
 'Westhus': ['Eric']
}

assert (list_to_dict() == answer)
```

## #2 Join using a dictionary

We have a list of patients, diagnosis, and length of stay.  We also have a dictionary that contains diagnosis and average length of stay.  Produce an output list that lists the patient and an indicator if the patient's stay was 'too long', 'too short', 'just right'

<!-- grade_id: cell-f3813e33810dc898, points: 0, grade: false, solution: true -->
```python
def patient_los(avg_los, patients):

    conclusions = []
    
    # YOUR CODE HERE
    for patient in patients:
        if patient[2] > avg_los[patient[1]]:
            conclusions.append([patient[0],"too long"])
        if patient[2] < avg_los[patient[1]]:
            conclusions.append([patient[0],"too short"])
        elif patient[2] == avg_los[patient[1]]:
            conclusions.append([patient[0],"just right"])
    
    return conclusions
```

```python
avg_los = {
    "Hemolytic jaundice and perinatal jaundice" : 2,
    "Medical examination/evaluation" : 3.2,
    "Liveborn" : 3.2,
    "Trauma to perineum and vulva" : 2.1,
    "Normal pregnancy and/or delivery" : 2,
    "Umbilical cord complication" : 2.1,
    "Forceps delivery" : 2.2,
    "Administrative/social admission" : 4.2,
    "Prolonged pregnancy" : 2.4,
    "Other complications of pregnancy" : 2.5
}

patients = [
    ['Boal', 'Medical examination/evaluation', 1.1],
    ['Boal', 'Other complications of pregnancy', 3.3],
    ['Jones', 'Liveborn', 3.2],
    ['Ashbury', 'Forceps delivery', 2.0]
]

patient_los(avg_los, patients)
```

Output:
```
[['Boal', 'too short'],
 ['Boal', 'too long'],
 ['Jones', 'just right'],
 ['Ashbury', 'too short']]
```

<!-- grade_id: cell-2af0a520bb56ee21, points: 10, grade: true, solution: false -->
```python
conclusions = [['Boal', 'too short'],
    ['Boal', 'too long'],
    ['Jones', 'just right'],
    ['Ashbury', 'too short']]

assert(patient_los(avg_los, patients) == conclusions)
```

```python

```