Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - Which of the following are ways of encapsulating code in Python?

> A) dict
>
> B) str
>
> C) function
>
> D) A and C
>
> E) All of the above

<!-- grade_id: cell-d789a32349f1885c, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "E"
    
    return answer
```

<!-- grade_id: cell-ebb0bb10a19c592a, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - Describe in your own words why abstraction is important in programming?


<!-- grade_id: cell-7ec686102170454b, points: 1, grade: true, solution: true -->
```python
def q02():

    # YOUR CODE HERE
    answer = "Abstraction is another way to reduce hard-coding. Once you have encapsulated code that you'll likely use multiple times, you can change it to be more generalizable to a wider scope of inputs. For example, if code has been encapsulated into a function, you can allow multiple types of inputs (parameters) for the function rather than requiring a specific type, or just using the original input you needed the code for."

    return answer
```

```python
assert type(q02()) == str
```

## #3 - What does DRY stand for?


> A) Do not repeat yourself
>
> B) Do not repeat yourself
>
> C) Do not repeat yourself
>
> D) All of the above

<!-- grade_id: cell-46ee97a89206f3d4, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "D"
    
    return answer
```

<!-- grade_id: cell-be350b2fc623fd53, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```