Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - When Pandas wants to display that it has no data for a particular record and column, which of the following is displayed?

> A) `None`
>
> B) `NaN`
>
> C) ` ` - blank
>
> D) `#NA`


<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: q1-solution, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - What of the following is the syntax for replacing all missing values in a particular column with a default value of `'Unknown'`?

> A) `data['col'].fillna('Unknown')`
>
> B) `fillna(data,'col','Unknown')`
>
> C) `data['col'].ifisna('Unknown')`
>
> D) `data['col'].isna() = 'Unknown'`


<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: q2-solution, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - Which one of the following will ensure that all of the values in `data['col']` are in lower case?

> A) `data['col'].str.lower()`
>
> B) `data.col = data.col.str.lower()`
>
> C) `data['col'].str.lower(data['col'])`
>
> D) `data['col'] = data['col'].lower()`


<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: q3-solution, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```