Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - How many times will the following code print the phrase 'Hello World!'?

```python
start = 1
stop = 1000

for num in range(start, stop):
    print('Hello World!')
    start = start + 5
```

> A) 200 times
>
> B) 199 times
>
> C) 999 times
>
> D) 201 times

<!-- grade_id: cell-d789a32349f1885c, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "C"
    
    return answer
```

<!-- grade_id: cell-ebb0bb10a19c592a, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - Triangular Number Sequence

There is a special sequence of numbers known as a triangular sequence.  It is generated by starting with `1` and adding `2` to that. Then adding `3` to the result.  Then adding `4` to the result of that.  The sequence continues by incrementally adding `1` more than the previous turn.  The first few numbers are shown below:

`1, 3, 6, 10, 15, ...`

Fill in the function below to find the `n`th number in that sequence.  Consider the answer `1` above to be the result of calling `triangle(1)` and `3` to be the result of `triangle(2)`, etc.  If `n` is less than 1, return 0.

<!-- grade_id: cell-7ec686102170454b, points: 0, grade: false, solution: true -->
```python
def triangle(n):
    
    # YOUR CODE HERE
    increment = 1
    answer = 0
    if n < 1:
        return 0
    else:
        for num in range(n):
            answer = answer + increment
            increment = increment + 1
    
        return answer
```

```python
# Feel free to test your code here
triangle(6)
```

Output:
```
21
```

<!-- grade_id: cell-0592dde608b988be, points: 1, grade: true, solution: false -->
```python
assert(triangle(1) == 1)
assert(triangle(2) == 3)
assert(triangle(3) == 6)
assert(triangle(4) == 10)
```

## #3 - Which of these best describes how the `range` function works?

> A) It generates a list of values that the for loop can then iterate over.
>
> B) It pre-generates the code required to execute the for loop according to the range() parameters
>
> C) If just runs the specified code as many times as necessary to create the right output.
>
> D) It returns an iterable object that returns the next value for each time through the loop.


<!-- grade_id: cell-46ee97a89206f3d4, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "D"
    
    return answer
```

<!-- grade_id: cell-be350b2fc623fd53, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```