Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 Variables
In the scenario below, which statments about the variables `a` and `b` are true.
```
a = [1, 2, 3]
b = a
```

> A) `a` is a list and `b` is a reference to another variables
> 
> B) `a` and `b` both refer to a list that has three elements
> 
> C) `b` is a copy of `a`
>
> D) `b` refers to `a` which refers to a list with three elements


<!-- grade_id: cell-a63a2c3bbc991423, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: cell-5f97db48f3a84609, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 Variable Names

Which of the following is **NOT** true about how you should name variables?

> A) Variable names should be as short as possible
> 
> B) Variable names can use abbreviations
>
> C) Variable names should be readable by humans
>
> D) Variable names represent a kind of documentation

<!-- grade_id: cell-ad5aa6c5e74c36be, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "D"
    
    return answer
```

<!-- grade_id: cell-de15d4a2d5d9f1a8, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 Why Use Variables

Based on the purpose and best practices for using variables in Python, which of the following lines of code represents the best use of variables?  The code in the example below is used to convert degrees Farenheit to degress Celsius.

> A) `v2 = (v1 - 32) * 5/9`
>
> D) `temp_c = (temp_f - 32) * 5/9`
>
> B) `tempcelsius = (tempfarenheit - 32) * 5/9`
>
> C) `temp-celsius = (temp-farenheight - 32) * 5/9`


<!-- grade_id: cell-d99e68f47741b529, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "D"
    
    return answer
```

<!-- grade_id: cell-eb1647b5986682c1, points: 0, grade: true, solution: false -->
```python
assert type(q03()) == str
```