Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - What will be the output of the command below?  (Don't run it, just figure it out)

```
date='January 14, 2017'
print(date[-4:])
```

> A) `Janu`
>
> B) Nothing.  It will raise an exception.
>
> C) Nothing.  Just and empty string.
>
> D) `2017`


<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "D"
    
    return answer
```

<!-- grade_id: q1-solution, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - Which code below will successfully return `E.`?

```
name = 'Paul E. Boal'
```

> A) `name.split(' ')[2]`
>
> B) `name[name.find(' ')+1:name.find(' ')+3]`
>
> C) `name[6:7]`
>
> D) `middle(name)`


<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: q2-solution, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - Use the ASCII chart from the slides or find another online to decode the message (in decimal values) below.

```
080 121 116 104 111 110 080 097 110 100 097 115
```


> A) Pandas String
>
> B) Python Pandas
>
> C) Str Expression
>
> D) Substrings

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: q3-solution, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```

```python

```