Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - Which of the following is the professor's preferred way to write dates?

> A) Monday, January 4, 2018
>
> B) 1/4/18
>
> C) 2018-01-04
>
> D) 1/4/2018


<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "C"
    
    return answer
```

<!-- grade_id: q1-solution, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - What kind of result does Pandas give when subtracting two dates?

> A) A `float` number of days
>
> B) An `offset` data type that contains components for days, hours, minutes, etc
>
> C) Another `datetime` whose year, month, day, etc are the number of those things between the start and end
>
> D) Nothing, this is not a valid operation


<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: q2-solution, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - Which one of the following does Pandas not have a built-in date offset for?

> A) semi-monthly
>
> B) quarterly
>
> C) fortnightly
>
> D) quarterly


<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "C"
    
    return answer
```

<!-- grade_id: q3-solution, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```