Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Week 2 Exercises

_McKinney 2.3_

**In the exercises below, don't worry about the `def` and other things around the code you need to write.  If there is a skeleton of code already there, just put your code where it says "your code here"**

## #1 Averaging three numbers

Image that you've got a series of body temperature measurements and we want to find the average of three of them.  Write code below that computes the average of three numbers represented by the variables `temp1`, `temp2`, `temp3` and assigns the result to the variables `avg_temp`.

<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
def average_temp(temp1, temp2, temp3):
    
    # YOUR CODE HERE
    avg_temp = (temp1 + temp2 + temp3) / 3
    
    return avg_temp
```

Your function should return 98 if you give it inputs of 96, 99, 99.  Check to see that it does by typeing `average_temp(96, 99, 101)` in the cell below and running it.

```python
average_temp(96, 99, 101)
```

Output:
```
98.66666666666667
```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
assert average_temp(96, 99, 99) == 98
assert average_temp(0, 0, 0) == 0
assert average_temp(100, -100, 0) == 0
assert average_temp(96, 98, 100) == 98
assert average_temp(96.5, 98.5, 100.5) == 98.5
assert average_temp(96, 98, 99) == (293 / 3)
```

## #2 Computing BMI

Create new names for the variables that you need in the code framework below and use the variables to compute the corresponding body mass index (BMI). You can use any meaningful variable names you like. When you compute the BMI, assign a variable to it, and then print the result to the notebook's output.

* Weght in kilograms is 90.7
* Height in centimeters in 182

Note that a simple BMI calculation is weight in kilograms divided by the square of height in meters.

For this problem, you'll be graded both on proper calculation as well as on programming style and appropriate variables names.

_Also note that BMI should not be considered a universal measure of health or fitness. A health BMI can vary greatly depending on many other fitness conditions and body types._

<!-- grade_id: bmi-answer, points: 0, grade: false, solution: true -->
```python
def bmi(parameter1, parameter2):
    
    # 1. Change the names of parameter1 and parameter2 to be something more meaningful
    # 2. Change the name of the return variable below
    
    wt = parameter1
    ht = parameter2
    
    #Needed to convert cm to m.
    ht = ht / 100
    
    bmi = wt / (ht**2)
    
    return bmi
    
    return return_value
```

Test your code below.  The BMI for a weight of 90.7 kg and height of 182 cm should be approximately 27.38

```python
bmi(90.7, 182)
```

Output:
```
27.38195870064002
```

<!-- grade_id: bmi-test, points: 0, grade: true, solution: false -->
```python
assert round(bmi(90.7, 182), 2) == 27.38
assert round(bmi(91, 180), 2) == 28.09
assert round(bmi(70, 170), 2) == 24.22
```

<!-- grade_id: bmi-style, points: 5, grade: true, solution: true -->
```python

```

## #3 Weight Loss

Assuming that someone doesn't change height over time, what is the approximate number of days it will take for someone to get from the BMI you found in #3 above to a BMI of 24.9 (typically regarded as being in a healthy range) while losing weight at a rate of 0.75 kg per week.

As needed, use variables and multiple calculation steps to make your thought process visible in the code. For instance, rather than using the notebook as a calculator, use variables for intermediate answers.

Again, you will receive points both for correct answers as well as for style.


<!-- grade_id: weight-answer, points: 0, grade: false, solution: true -->
```python
def lose_weight(weight, height, kg_per_week, target_bmi):
    
    # Note that the variables you need area already named for you:
    # height is the height in cm
    # weight is the weight in kg
    # kg_per_week is how much the person wants to lose per week
    # target_bmi is the intended target bmi
    
    # YOUR CODE HERE
    height = height / 100
    target_wt = target_bmi * (height**2)
    days = ((weight - target_wt) / kg_per_week) * 7

    return days
```

Test your code below.  It should return approximately 76.73 days for the given parameters above.

```python
lose_weight(90.7, 182, 0.75, 24.9)
```

Output:
```
76.73157333333342
```

<!-- grade_id: weight-tests, points: 5, grade: true, solution: false -->
```python
assert round(lose_weight(90.7, 182, 0.75, 24.9),2) == 76.73
assert round(lose_weight(91, 180, 0.75, 24.9),2) == 96.36
assert round(lose_weight(70, 170, 0.75, 22),2) == 59.92
assert round(lose_weight(70, 170, 2, 22),2) == 22.47

```

<!-- grade_id: weight-style, points: 5, grade: true, solution: true -->
```python

```