Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - Select all of the data types that can be represented within a JSON format:

Provide your answer in the form of a list.  For example: `answer = ["A","B"]` 

> A) str
>
> B) list
>
> C) dict
>
> D) int


<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = ["A","D"]
    
    return answer
```

<!-- grade_id: q1-solution, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == list
```

## #2 - JSON stands for:

> A) Just a Simple Object Notation
>
> B) JavaScript Object Notation
>
> C) JavaScript Objects and Numbers
>
> D) Java Standard Object Notation


<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: q2-solution, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - Without just running the code, what will the following code print?

```
import json

data = """
[  {"id": { "last name": "Boal", "first name": "Paul", "middle initial": "E" }},
   {"id": { "last name": "Westhus", "first name": "Eric" }} ]
"""

names = json.loads(data)
list(names[1]['id'].keys())[1]
```

> A) 'Eric'
>
> B) 'first name'
>
> C) 'last name'
>
> D) 'Boal'


<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: q3-solution, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```