Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - Which statment about the following code is NOT correct if you were to run it with our Jupyter site?

```python
import requests
response = requests.get('http://google.com')
```

---

> A) A Python process on jupyter.slucor.net is requesting a web page from Google
>
> B) Google is sending information to jupyter.slucor.net
>
> C) Google will send the HTML contents of its default page to a Python process
>
> D) Google is sending information to your personal laptop


<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "D"
    
    return answer
```

<!-- grade_id: q1-solution, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - What kind(s) of files can you find on the Internet?

Choose all that apply and answer using a list (e.g. `answer = ["A", "C", "D"]`)

---

> A) CSV Files
>
> B) Images
>
> C) HTML
>
> D) JSON


<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = ["A","B","C","D"]
    
    return answer
```

<!-- grade_id: q2-solution, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == list
```

## #3 - Which kind of information on the Internet do you think will be hardest to retrieve and turn into something usable?

> A) CSV files from data.gov
>
> B) Data embedded within an HTML file
>
> C) JSON documents from a web service
>
> D) Structure XML documents


<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: q3-solution, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```