Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Week 2 Exercises

_McKinney 2.3_

**In the exercises below, don't worry about the `def` and other things around the code you need to write.  If there is a skeleton of code already there, just put your code where it says "your code here"**

## #1 NAND

Inside of computer chips, it's often easier to do a logical operation called NAND (not and).  Write a function that takes two boolean parameters and returns the NAND of those values.  Here is the k-map for NAND:

```
NAND: 

        TRUE   FALSE
TRUE    False  True
FALSE   True   True
```

Include the four tests from the k-map in your docstring, and show the output of testing your function for correct operation using testmod() as demonstrated during class.

<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
def nand(a, b):
    
    # YOUR CODE HERE
    answer = not (a and b) 
    
    return answer
```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
assert nand(True, True) == False
assert nand(True, False) == True
assert nand(False, True) == True
assert nand(False, False) == True
```

## #2 LACE Score
There is a simple readmission index called the LACE Score: https://www.hindawi.com/journals/bmri/2015/169870/tab1/

Use this documentation to create a function that can compute LACE scores based on the 4 input parameters.  Below is the function signature and documentation that you start with.

<!-- grade_id: bmi-answer, points: 0, grade: false, solution: true -->
```python
def LACE(length_of_stay, acute_flag, charlson, ed_visits):
    """(int, bool, int, int) -> int
    This function uses the logic from https://www.hindawi.com/journals/bmri/2015/169870/tab1/
    to compute the LACE score for this patient.
    
    >>> LACE(4, False, 1, 0)
    5
    
    >>> LACE(4, True, 4, 7)
    16
    
    """
    
    # YOUR CODE HERE
    score = 0
    if length_of_stay >= 1:
        if length_of_stay >= 14:
            score = score + 7
        elif 7 <= length_of_stay <= 13:
            score = score + 5
        elif 4 <= length_of_stay <= 6:
            score = score + 4
        elif length_of_stay == 3:
            score = score + 3
        elif length_of_stay == 2:
            score = score + 2
    else:
        score = score + 1
    if acute_flag:
        score = score + 3
    if charlson >=1:
        if charlson >= 4:
            score = score + 5
        elif charlson == 3:
            score = score + 3
        elif charlson == 2:
            score = score + 2
    else:
        score = score + 1
    if ed_visits >= 1:
        if ed_visits >= 4:
            score = score + 4
        if ed_visits == 3:
            score = score + 3
        if ed_visits == 2:
            score = score + 2
    else:
        score = score + 1
    
    return score
```

```python
LACE(4, True, 4, 7)
```

Output:
```
16
```

<!-- grade_id: bmi-test, points: 10, grade: true, solution: false -->
```python
assert LACE(4, False, 1, 0) == 5
assert LACE(4, True, 4, 7) == 16
```

## #3 Care Management Criteria

Care managers use LACE as part of the criteria for assigning a care coodinator to a patient who has been recently discharged. If the score is above 10, then a care coordinator will be assigned. The other criteria they use is if the patient has been discharged with a diagnosis of CHF or COPD.  If the diagnosis field has CHF or COPD in it, then the patient will have a care coordinator assigned.

For this exercise, write another function that takes the same inputs as LACE() plus another diagnosis parameter, and return True or False depending on if the paient needs a care coordinator.

**NOTE** Pay attention to the fact that the order of parameters in this function definition are not the same as the order they were in the LACE score.  

<!-- grade_id: weight-answer, points: 0, grade: false, solution: true -->
```python
def assign_care_coordinator(diagnosis_cd, ed_visits, length_of_stay, acute_flag, charlson):
    """ (str, int, int, bool, int) -> bool
    Care managers use LACE as part of the criteria for assigning a care coodinator to a 
    patient who has been recently discharged. If the score is above 10, then a care 
    coordinator will be assigned. The other criteria they use is if the patient has been 
    discharged with a diagnosis of CHF or COPD. If the diagnosis field has CHF or COPD in 
    it, then the patient will have a care coordinator assigned.
    
    >>> assign_care_coordinator('None', 0, 4, False, 1)
    False

    >>> assign_care_coordinator('CHF', 0, 4, False, 1)
    True

    >>> assign_care_coordinator('COPD', 0, 4, False, 1)
    True

    >>> assign_care_coordinator('None', 7, 4, True, 4)
    True
    
    >>> assign_care_coordinator('CHF', 7, 4, True, 4)
    True
    """
    
    # YOUR CODE HERE
    score = 0
    assignment = False
    if length_of_stay >= 1:
        if length_of_stay >= 14:
            score = score + 7
        elif 7 <= length_of_stay <= 13:
            score = score + 5
        elif 4 <= length_of_stay <= 6:
            score = score + 4
        elif length_of_stay == 3:
            score = score + 3
        elif length_of_stay == 2:
            score = score + 2
    else:
        score = score + 1
    if acute_flag:
        score = score + 3
    if charlson >=1:
        if charlson >= 4:
            score = score + 5
        elif charlson == 3:
            score = score + 3
        elif charlson == 2:
            score = score + 2
    else:
        score = score + 1
    if ed_visits >= 1:
        if ed_visits >= 4:
            score = score + 4
        if ed_visits == 3:
            score = score + 3
        if ed_visits == 2:
            score = score + 2
    else:
        score = score + 1
    
    if score > 10 or diagnosis_cd != "None":
        assignment = True
        
    return assignment
```

<!-- grade_id: weight-tests, points: 5, grade: true, solution: false -->
```python
assert assign_care_coordinator('None', 0, 4, False, 1) == False
assert assign_care_coordinator('CHF', 0, 4, False, 1) == True
assert assign_care_coordinator('COPD', 0, 4, False, 1) == True
assert assign_care_coordinator('None', 7, 4, True, 4) == True
assert assign_care_coordinator('CHF', 7, 4, True, 4) == True
```

```python

```