Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - Which of the following is least likely to be used as a delimiter in structured files?

> A) `,`
>
> B) `|`
>
> C) `[tab]`
>
> D) `[space]`


<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "D"
    
    return answer
```

<!-- grade_id: q1-solution, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - Which approaches can be used to read pipe-delimited (`|`) files?

*Report your answer as a list of options, e.g.: `["A","C"]`*

> A) Using the csv module
>
> B) Using the pandas module
>
> C) Using built-in file reading
>
> D) Using the psv module


<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = ["A","B"]
    
    return answer
```

<!-- grade_id: q2-solution, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == list
```

## #3 - In a CSV file, how does the format change if you want to have a `,` in the actual contents of the file rather than as a delimiter?

> A) You can't do that.  You have to remove every occurence of `,`s
>
> B) Enclose the `,` using quotation marks
>
> C) Escape the `,` using a back-slash
>
> D) Nothing special, it will be obvious to the reader


<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: q3-solution, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```