Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

## #1 List of Allergies

In this directory, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
import json

# YOUR CODE HERE
def allergy_count(json_file):
    """
    This function first reads in the specified json_file.
    It then creates a counter for every entry of 'substance' found in the 'entry' subset of the file.
    Each additional 'substance' found adds 1 to the counter, as there is only one item for each 'substance' key.
    """
    with open(json_file) as i:
        allergy = json.load(i)
    
    items = 0
    for case in allergy['entry']:
        if 'substance' in case['resource']:
            items += 1
    
    return items
```

```python
allergy_count('allergies.json')
```

Output:
```
3
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count('allergies.json')) == int

```

## #2 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
import json

# YOUR CODE HERE
def allergy_list(json_file):
    """
    I think I might need more practice subsetting nested dictionaries, nonetheless:
    This function will grab the patient name, allergic substance, and allergic reaction
    for every instance in 'entry' of the json_file provided and combine them into a tidy
    list of lists.
    """
    with open(json_file) as i:
        allergy = json.load(i)
    
    list1 = []
    allergy_l = []
    for case in allergy['entry']:
        list1 = ([case['resource']['patient']['display'],case['resource']['substance']['text'], case['resource']['reaction'][0]['manifestation'][0]['text']])
        allergy_l.append(list1)
    
    return allergy_l
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
output=[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]

assert allergy_list('allergies.json') == output

```

## #2 Allergy Reaction

Write a function called `allergy_reaction(json_file,substance)` that takes two parameter and returns the reaction that will happen if the patient takes the specified substance.  Solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
import json

# YOUR CODE HERE
def allergy_reaction(json_file,substance):
    """
    This function searches for a keyword 'substance' in the list of lists generated
    by running allergy_list on a specified json_file. It returns the reaction associated
    with an allergic reaction to the substance provided. If there are no matches, nothing
    will be returned.
    """
#I'm unsure how you will react to this. My lack of manually returning None wasn't spurred by
#laziness, rather I was playing with the function and it just sort of ended up working this way.
#However I can see why it wouldn't be a very 'safe' function. :)

    allergies = allergy_list(json_file)
    for lst in allergies:
        if substance in lst:
            return lst[2]
```

```python
import json

def allergy_reaction(json_file,substance):
    """
    This function searches for a keyword 'substance' in the list of lists generated
    by running allergy_list on a specified json_file. It returns the reaction associated
    with an allergic reaction to the substance provided. If there are no matches, nothing
    will be returned.
    """
#I gave it a second try to see if I could at least find a simple way to accomplish manually
#having None be the output if there weren't any matches. and I succeeded! So here is v1.2 of allergy_list.
    allergies = allergy_list(json_file)
    for lst in allergies:
        if substance in lst:
            return lst[2]
            break
            return None
```

```python
import json

allergy_reaction('allergies.json','PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction('allergies.json','PENICILLIN G') == 'Hives'
assert allergy_reaction('allergies.json','SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction('allergies.json','STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction('allergies.json','PENICILLIN') == None
```