Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - When Python tells you there's been an exception, what is the term we use for that?

> A) Throw
>
> B) Raise
>
> C) Handle
>
> D) Crash

<!-- grade_id: cell-d789a32349f1885c, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: cell-ebb0bb10a19c592a, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - Explain in your own words why you would want to catch and exception in your code?


<!-- grade_id: cell-7ec686102170454b, points: 1, grade: true, solution: true -->
```python
def q02():

# YOUR CODE HERE
    answer = "If Python didn't raise an exception/give an error when it was incapable of doing what you ask it, it would be insanely difficult to figure out why Python wasn't giving an output, or wasn't giving an output you expected."
    
    return answer
```

```python
assert type (q02()) == str
```

## #3 - Explain in your own words why simply returning -1 or None or some other default value might not be a good idea if your program encounters an error?


<!-- grade_id: cell-46ee97a89206f3d4, points: 1, grade: true, solution: true -->
```python
def q03():

# YOUR CODE HERE
    answer = "It could be bad for a few reasons. One in particular would be if someone else is using your function. Python isn't raising an error but is returning a value of either -1 or None, which doesn't intrinsically indicate that an error occurred. There are also cases where a result of -1 or None would be appropriate, depending on the function. This means the result could be the wrong output, but the user wouldn't necessarily know since there wasn't a message output about what went wrong, or at least a warning about what could have happened."

    return answer
```

```python
assert type (q03()) == str
```