Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - An IF statement in Python can:

> A) Only perform one conditional test
>
> B) Returns the indented actions if the test is True, or None if it is not True
>
> C) Can contain as many branches of logic as the programmer wants
>
> D) Executes the next line of code only if the test is True, then moves to always execute the remaining code

<!-- grade_id: cell-d789a32349f1885c, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "C"
    
    return answer
```

<!-- grade_id: cell-ebb0bb10a19c592a, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - What will the following code print to the screen?

```python
age = 21

if age > 18:
    label = 'adult'
label = 'kid'

print(label)
```

> A) `kid`
>
> B) `adult`
>
> C) `adult` (new line) `kid`
>
> D) _nothing_

<!-- grade_id: cell-7ec686102170454b, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: cell-0592dde608b988be, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - What will the follow code print to the screen?

```python
temp = 22
precipitation = 'wet'

if precipitation != 'wet':
    weather = 'dry'
else:
    if temp > 32:
        weather = 'rain'
    elif temp < 32:
        weather = 'sleet'
    elif temp < 28:
        weather = 'mix'
    else:
        weather = 'snow'

print(weather)
```

> A) `sleet`
>
> B) `mix`
>
> C) `now`
>
> D) `dry'

<!-- grade_id: cell-46ee97a89206f3d4, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = "A"
    
    return answer
```

<!-- grade_id: cell-be350b2fc623fd53, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == str
```