Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emily Holcomb"
COLLABORATORS = ""
```

---

# Instructions

In the answer section below, select your answer to the following multiple choice questions by typing `answer = "X"` where `X` is the letter for the correct multiple choice answer.

## #1 - Which one of the following is not a valid syntax for creating a filter for a data frame to only those records where the `state` is `CA` ?  Below is one example that is valid.  Choose the option that does NOT achieve the same result.

```
data['state'] == 'CA'
```

> A) `data[state == 'CA']`
>
> B) `data.state == 'CA'`
> 
> C) `data['state'].str.contains('^CA$')`
>
> D) `data['state'].apply(lambda x: x == 'CA')`


<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
def q01():
    
    # YOUR CODE HERE
    answer = "B"
    
    return answer
```

<!-- grade_id: q1-solution, points: 1, grade: true, solution: false -->
```python
assert type(q01()) == str
```

## #2 - What does the following code do, assuming that `data` had a column called `phone`?

```
filter = data['phone'].str.contains("^[0-9]{3}.[0-9]{3}.[0-9]{4}$")
results = data[~filter]
results.shape
```

> A) Filters the `data` for records that have only numerical digits in the phone field
>
> B) Filters the `data` for records that have a valid US phone number format of `### ### ####`
>
> C) Filters the `data` for records that do not have a valid US phone number format of `###-###-####`
>
> D) Filters the `data` for records that do not have a valid US phone number format of `###x###x####` where `x` can be any character


<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def q02():
    
    # YOUR CODE HERE
    answer = "C"
    
    return answer
```

<!-- grade_id: q2-solution, points: 1, grade: true, solution: false -->
```python
assert type(q02()) == str
```

## #3 - Which of the following option(s) can be used to select information from a data frame?

> A) SQL, using pandasql
>
> B) Regular Expressions
>
> C) Boolean Logic
>
> D) Any arbitrary function you write


<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def q03():
    
    # YOUR CODE HERE
    answer = ["A","B","C"]
    
    return answer
```

<!-- grade_id: q3-solution, points: 1, grade: true, solution: false -->
```python
assert type(q03()) == list
```